Перед началом прохождения убедитесь, что все настроено как надо. Во-первых, **перезапустите ядро** (**restart the kernel**) (в панели меню выберите Kernel$\rightarrow$Restart) и затем **запуск всех ячеек** (**run all cells**) (в панели меню выберите Cell$\rightarrow$Run All).

Убедитесь, что вы заменили всем такие места 
```python
# YOUR CODE HERE
raise NotImplementedError()
```
и
```
"YOUR ANSWER HERE"
```
на свои решения.

Также, заполните свои данные (имя и группу):

```python
NAME = "Вячеслав Завьялов"
GROUP = "6492"
```

```python
import uuid

STUDENT_UUID = uuid.uuid4()
print(f"Your uuid is {STUDENT_UUID}")
```

Output:
```
Your uuid is cd385d0f-bb2a-4994-8ed8-634151764ce7
```

---

# Приветствие

Привет, в этом тесте вы познакомитесь с тремя примерами задач, которые будут у вас в будущих контрольных точках. 

Задачи не сложные и нацелены лишь на то, чтобы просто вы опробовали свои силы и познакомились с форматом!

## Задание 1 (1 балл)

Напишите код разворачивания списка.

Вместо блока
```python
# YOUR CODE HERE
raise NotImplementedError()
```
нужно реализовать свой код, чтобы он проходил по следующим тестам.

Сразу учитывайте, что не все тесты открыты и часть из них будет проверяться после сдачи решения.

<!-- grade_id: cell-b7df227e69ae6723, points: 0, grade: false, solution: true -->
```python
def reverse_list(data: list):
    data = data[::-1]
    return(data)
```

<!-- grade_id: cell-a9075953f720b87b, points: 1, grade: true, solution: false -->
```python
assert reverse_list([0]) == [0]
assert reverse_list([3, 2, 1]) == [1, 2, 3]
assert reverse_list([]) == []

```

## Задание 2 (1 балл)

Разработайте функцию вычисления площади круга по переданному радиусу. В случае переданного отрицательного радиуса должен вернуться 0.

<!-- grade_id: cell-e0913548aff9db16, points: 0, grade: false, solution: true -->
```python
import math as m

def compute_circle_area(radius: float):
    if radius >= 0:
        return(m.pi*radius**2)
    else:
        return 0
```

<!-- grade_id: cell-4302ce285797495b, points: 1, grade: true, solution: false -->
```python
assert m.isclose(compute_circle_area(3), 28.274, abs_tol=1e-3)
assert m.isclose(compute_circle_area(-2), 0)

```

## Задание 3 (2 балла)

Ответьте кратко на вопрос.
Какие бывают типы данных Immutable, а какие Mutable? В чем разница между Immutable и Mutable типами?

<!-- grade_id: cell-5008a206312d2427, points: 2, grade: true, solution: true -->
Неизменяемые типы после создания не могут быть изменены, а изменяемые можно менять на протяжении всей программы. 
Изменяемыми являются list, dict, set, bytearray
Неизменяемыми являются bool, None, int, float, complex, str, bytes, tuple, range, fronzeset