Перед началом прохождения убедитесь, что все настроено как надо. Во-первых, **перезапустите ядро** (**restart the kernel**) (в панели меню выберите Kernel$\rightarrow$Restart) и затем **запуск всех ячеек** (**run all cells**) (в панели меню выберите Cell$\rightarrow$Run All).

Убедитесь, что вы заменили всем такие места 
```python
# YOUR CODE HERE
raise NotImplementedError()
```
и
```
"YOUR ANSWER HERE"
```
на свои решения.

> Не переименовывайте ноутбук!

Также, заполните свои данные (имя и группу):

```python
NAME = "Вячеслав"
GROUP = "6492"
```

```python
import uuid

STUDENT_UUID = uuid.uuid4()
print(f"Your uuid is {STUDENT_UUID}")
```

Output:
```
Your uuid is be895489-7177-4470-8aff-0f4aada0fddb
```

---

# Приветствуем тебя, путник!

Это первая контрольная работа с некоторыми вопросами по теме numpy, pandas и python.




```python
import numpy as np
import pandas as pd

np.random.seed(123)
```

### Задание 1 (1 балл)

Напишите код для разворачивания строки и приведения её к нижнему регистру. 

Например, на вход подаётся строка "Привет", на выходе - "тевирп".

<!-- grade_id: cell-69a75a6a9f247b59, points: 0, grade: false, solution: true -->
```python
def reverse_string(string: str) -> str:
    # YOUR CODE HERE
    string = string[::-1].lower()
    return string
```

<!-- grade_id: cell-6c2c7d17e93c9be4, points: 1, grade: true, solution: false -->
```python
assert len(reverse_string("hi")) == 2 
assert reverse_string("H") == "h"
assert reverse_string("") == ""

```

### Задание 2 (1 балл)

Выведите уникальные строки из входного массива (матрицы).

<!-- grade_id: cell-a83f3275964bb308, points: 0, grade: false, solution: true -->
```python
def get_unique_rows(matrix: np.ndarray) -> np.ndarray:
    # YOUR CODE HERE
    matrix = np.unique(matrix,axis=0)
    return matrix
```

<!-- grade_id: cell-4d50a5c607725df3, points: 1, grade: true, solution: false -->
```python
assert np.allclose(get_unique_rows(np.array([[1]])), [1])
assert np.allclose(get_unique_rows(np.array([[1], [1], [1]])), [1])
assert np.allclose(get_unique_rows(np.array([[1, 2], [1, 2], [2, 1]])), [[1, 2], [2, 1]])

```

### Задание 3 (1 балл)

Выведите N наибольших значений из массива (матрицы). 

> Если N больше, чем общее число элементов в массиве, то нужно вернуть весь исходный массив. 

<!-- grade_id: cell-d4aba3080a3cd06f, points: 0, grade: false, solution: true -->
```python
def get_n_largest(arr: np.ndarray, n: int) -> np.ndarray:
    # YOUR CODE HERE
    arr = np.sort(arr)
    arr = arr[::-1]
    if n > len(arr):
        n = len(arr)
    a = np.zeros(n)
    for i in range(n):
        a[i] = arr[i]
    print(a)
    return a
```

<!-- grade_id: cell-922af51c5bbde513, points: 1, grade: true, solution: false -->
```python
assert len(get_n_largest(np.array([2, 1, 3]), 1)) == 1
assert np.allclose(get_n_largest(np.zeros(5), 6), np.zeros(5))

```

Output:
```
[3.]
[0. 0. 0. 0. 0.]
```

### Задание 4 (1 балл)

Найдите самое часто повторяющееся число в массиве. 

<!-- grade_id: cell-db76e26eedba62d1, points: 0, grade: false, solution: true -->
```python
def get_freq_value(arr: np.ndarray) -> int:
    a = set(list(arr))
    arr=list(arr)
    max = None 
    max1 = 0
    for i in a:
        b = arr.count(i)
        if b > max1:
            max1 = b
            max = i 
    return int(max)
    
```

<!-- grade_id: cell-6617d7118be7f720, points: 0, grade: true, solution: false -->
```python
assert isinstance(get_freq_value(np.random.randint(0, 10, 10)), int)
assert get_freq_value(np.array([1, 1])) == 1
assert get_freq_value(np.array([1, 3, 2, 2])) == 2

```

### Задание 5 (1 балл)

Преобразуйте два списка в pandas.Series, где первый списoк - это индексы, а второй - значения. 

<!-- grade_id: cell-fa84a84ced474d2f, points: 0, grade: false, solution: true -->
```python
def get_series(list_1: list, list_2: list) -> pd.Series:
    # YOUR CODE HERE
    a = pd.Series(list_2, index = list_1)
    return a
```

<!-- grade_id: cell-110c1c0a2650e2fa, points: 1, grade: true, solution: false -->
```python
assert isinstance(get_series(["a"], [1]), pd.Series)
assert np.allclose(get_series(["a"], [1]).shape, 1)
assert all(get_series(["a", "c"], [1, 3]).index == ["a", "c"])
assert np.allclose(get_series(["a", "c", "1"], [1, 4, 5]).values, [1, 4, 5])

```

### Задание 6 (1 балл)

Разделите одну колонку со строками на 2 колонки: STD, City. 

<!-- grade_id: cell-58dc34e723ed3302, points: 0, grade: false, solution: true -->
```python
def separate_columns() -> pd.Series: 
    df = pd.DataFrame(
        [
            "STD, City",
            "33, Kolkata",
            "44, Chennai",
            "40, Hyderabad",
            "80, Bangalore"
        ], 
        columns=['row']
    )
    # YOUR CODE HERE
    d = df['row']
    columns = d[0].split(', ')
    dd = list(np.zeros(len(d)-1))
    for i in range(1,len(df)):
        dd[i-1] = d[i].split(', ')
    df = pd.DataFrame (
        data = dd[0:len(dd)],
        columns = columns
    )

    return(df)

```

<!-- grade_id: cell-ddbd80fb8adda273, points: 1, grade: true, solution: false -->
```python
assert set(separate_columns().columns) == set(["STD", "City"])
assert np.allclose(separate_columns().shape, (4, 2))

```

### Задание 7 (1 балл)

Выберите только те строки, в которых животное - это собака, а возраст меньше 3-х, не включая строки с NaN.

<!-- grade_id: cell-92c17c0ec05456d2, points: 0, grade: false, solution: true -->
```python
def selection():
    df = pd.DataFrame(
        zip(['cat', 'cat', 'dog', 'dog', 'dog'], [2.5, 3, 0.5, np.nan, 5]),
        columns=["animal", "age"]
    )
    # YOUR CODE HERE
    mask = ((df['animal']=='dog')&(df['age']<3))
    return df[mask]
    
```

<!-- grade_id: cell-2c583ee44d239300, points: 0, grade: true, solution: false -->
```python
assert np.allclose(selection().shape, (1, 2))
assert not (selection().animal.isnull()).all()

```

Output:
```
animal  age
0    cat  2.5
1    cat  3.0
2    dog  0.5
3    dog  NaN
4    dog  5.0
  animal  age
2    dog  0.5
  animal  age
0    cat  2.5
1    cat  3.0
2    dog  0.5
3    dog  NaN
4    dog  5.0
  animal  age
2    dog  0.5
```

### Задание 8 (2 балла)

Что означает `self` в языке Python? Зачем и где оно используется?

> Для примеров кода можно использовать конструкцию:

```python
def code_sample():
    pass
```

<!-- grade_id: cell-9547a4e31a45077b, points: 0, grade: true, solution: true -->
self - ссылка на определенный экземпляр в классе. Можно использовать и другое имя. То есть при вызове метода в него передается ссылка на объект и внутри метода можно обратиться к ней через self.

### Задание 9 (2 балла)

Опишите `broadcasting` механизм во фреймворке `numpy`. Приведите примеры.

<!-- grade_id: cell-e094babf8b515b35, points: 2, grade: true, solution: true -->
Это автоматическое расширение размерности. Например, np.array([0] * 3) выведет [0,0,0]

### Задание 10 (2 балла)

В чем разница между `.iloc`, `.loc` и `[]` операторами? Приведите примеры использования и объясните причины применения.

<!-- grade_id: cell-37ca0151f6ea53e2, points: 2, grade: true, solution: true -->
loc получает строки по меткам, iloc - по индексам. Если взять loc[1], то вернется столбец с меткой 1, если iloc[1], то строка с индексом 1